# Project Euler: Problem 8

https://projecteuler.net/problem=8

The four adjacent digits in the 1000-digit number that have the greatest product are 9 × 9 × 8 × 9 = 5832.

(see the number below)

Find the thirteen adjacent digits in the 1000-digit number that have the greatest product. What is the value of this product?

Use NumPy for this computation

```python
import numpy as np
```

```python
d1000 = "7316717653133062491922511967442657474235534919493496983520312774506326239578318016984801869478851843858615607891129494954595017379583319528532088055111254069874715852386305071569329096329522744304355766896648950445244523161731856403098711121722383113622298934233803081353362766142828064444866452387493035890729629049156044077239071381051585930796086670172427121883998797908792274921901699720888093776657273330010533678812202354218097512545405947522435258490771167055601360483958644670632441572215539753697817977846174064955149290862569321978468622482839722413756570560574902614079729686524145351004748216637048440319989000889524345065854122758866688116427171479924442928230863465674813919123162824586178664583591245665294765456828489128831426076900422421902267105562632111110937054421750694165896040807198403850962455444362981230987879927244284909188845801561660979191338754992005240636899125607176060588611646710940507754100225698315520005593572972571636269561882670428252483600823257530420752963450"

```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
#initalizes d1000 into an array, would want it to be a string but could not figure out how to properly
#match up indices. 

f = np.arange(0,1000)
for x in range(1000):
    f[x] = d1000[x]
#print (f)
# i = 13 because must be 13 numbers, so must start at 13
i = 13
b = []
d = []
#Loops through while i < the size of a - 13
while i < len(f):
    #Counting variable
    
    #Slices a into 13 numbers
    c = f[i-13:i]
    
    #Deals with when i would go out of the index
    if i == len(f):
        break
    #b = prod of 13 numbers
    b = np.cumprod(c)
    d.append([b[-1]])
    i = i + 1
#sorts d and then prints the largest number (the last index)
d.sort()
print (d[-1])
    

    #Initial thoughts. 
"""The way I wanted to do this problem was as follows:
I would convert d1000 to a string and then place that string into an array using the indices
of the string as the indices of the array. Doing this I would loop through each index multipying the following
12 indices by the current one. I would include an if statement that broke the loop when the index was
13 away from the end. (This way it would stay in the range) I would put all of these values into an array and 
sort that array and print out the last element in the array. Unfortunately, I couldn't put the string into the
array correctly. However, once I got over this hump, I believe my code would have worked. """
```

Output:
```
[23514624000]
"The way I wanted to do this problem was as follows:\nI would convert d1000 to a string and then place that string into an array using the indices\nof the string as the indices of the array. Doing this I would loop through each index multipying the following\n12 indices by the current one. I would include an if statement that broke the loop when the index was\n13 away from the end. (This way it would stay in the range) I would put all of these values into an array and \nsort that array and print out the last element in the array. Unfortunately, I couldn't put the string into the\narray correctly. However, once I got over this hump, I believe my code would have worked. "
```

<!-- grade_id: projecteuler8, points: 10, grade: true, solution: false -->
```python
assert True # leave this for grading
```

```python

```