# Project Euler: Problem 17

https://projecteuler.net/problem=17

If the numbers 1 to 5 are written out in words: one, two, three, four, five, then there are 3 + 3 + 5 + 4 + 4 = 19 letters used in total.

If all the numbers from 1 to 1000 (one thousand) inclusive were written out in words, how many letters would be used?


NOTE: Do not count spaces or hyphens. For example, 342 (three hundred and forty-two) contains 23 letters and 115 (one hundred and fifteen) contains 20 letters. The use of "and" when writing out numbers is in compliance with British usage.

First write a `number_to_words(n)` function that takes an integer `n` between 1 and 1000 inclusive and returns a list of words for the number as described above

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python



def number_to_words(n):
    """Given a number n between 1-1000 inclusive return a list of words for the number."""
    numbers = {1:"one", 2:"two", 3:"three", 4:"four",5:"five", 6:"six", 7:"seven", 8:"eight", 9:"nine",
    10:"ten", 11:"eleven", 12:"twelve", 13:"thirteen", 14:"fourteen", 15:"fifteen", 
    16:"sixteen", 17:"seventeen",18:"eighteen",19:"nineteen", 
    20:"twenty",30:"thirty", 40:"forty", 50:"fifty",60:"sixty",70:"seventy",80:"eighty",90:"ninety",
    100: "hundred", 1000:"thousand"}

    numStr = str(n)
    a = []
    p = len(numStr)
    
    s = n
    while s !=0:
        if s <=20:
            a.append(numbers[s])
        if 20<s<30:
            f = str(s)
            a.append(numbers[20]+numbers[int(f[1])])
        if 30<s<40:
            f = str(s)
            a.append(numbers[30]+numbers[int(f[1])])
        if 40<s<50:
            f = str(s)
            a.append(numbers[40]+numbers[int(f[1])])  
        if 50<s<60:
            f = str(s)
            a.append(numbers[50]+numbers[int(f[1])])
        if 60<s<70:
            f = str(s)
            a.append(numbers[60]+numbers[int(f[1])])
        if 70<s<80:
            f = str(s)
            a.append(numbers[70]+numbers[int(f[1])])
        if 80<s<90:
            f = str(s)
            a.append(numbers[80]+numbers[int(f[1])])
        if 90<s<100:
            f = str(s)
            a.append(numbers[90]+numbers[int(f[1])])
        if s == 100:
            a.append(numbers[1] + numbers[100])
        if s == 200:
            a.append(numbers[2] + numbers[100])

        if 100<s<1000:
            f = str(s)
            l = f[0]
            a.append(numbers[int(l)] + numbers[100]+"and" + numbers[int(f[2])])

            
            
        s = s-1
    return a
print (number_to_words(109))




```

Output:
```
['onehundredandnine', 'onehundredandeight', 'onehundredandseven', 'onehundredandsix', 'onehundredandfive', 'onehundredandfour', 'onehundredandthree', 'onehundredandtwo', 'onehundredandone', 'onehundred', 'ninetynine', 'ninetyeight', 'ninetyseven', 'ninetysix', 'ninetyfive', 'ninetyfour', 'ninetythree', 'ninetytwo', 'ninetyone', 'eightynine', 'eightyeight', 'eightyseven', 'eightysix', 'eightyfive', 'eightyfour', 'eightythree', 'eightytwo', 'eightyone', 'seventynine', 'seventyeight', 'seventyseven', 'seventysix', 'seventyfive', 'seventyfour', 'seventythree', 'seventytwo', 'seventyone', 'sixtynine', 'sixtyeight', 'sixtyseven', 'sixtysix', 'sixtyfive', 'sixtyfour', 'sixtythree', 'sixtytwo', 'sixtyone', 'fiftynine', 'fiftyeight', 'fiftyseven', 'fiftysix', 'fiftyfive', 'fiftyfour', 'fiftythree', 'fiftytwo', 'fiftyone', 'fortynine', 'fortyeight', 'fortyseven', 'fortysix', 'fortyfive', 'fortyfour', 'fortythree', 'fortytwo', 'fortyone', 'thirtynine', 'thirtyeight', 'thirtyseven', 'thirtysix', '
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: projecteuler17a, points: 4, grade: true, solution: false -->
```python
assert True # use this for grading the number_to_words tests.
```

Now define a `count_letters(n)` that returns the number of letters used to write out the words for all of the the numbers `1` to `n` inclusive.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def count_letters(n):
    """Count the number of letters used to write out the words for 1-n inclusive."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

Now write a set of `assert` tests for your `count_letters` function that verifies that it is working as expected.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: projecteuler17b, points: 4, grade: true, solution: false -->
```python
assert True # use this for grading the count_letters tests.
```

Finally used your `count_letters` function to solve the original question.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: projecteuler17c, points: 2, grade: true, solution: false -->
```python
assert True # use this for gradig the answer to the original question.
```